Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Потемкин Виктор"
COLLABORATORS = ""
```

---

Write a function which solves a quadratic equation, $x^2 + b x + c = 0$. Your function must receive two values for coefficients and return two real or complex roots. Always return two roots, even if the equation has a single double root. 

<!-- grade_id: cell-2032efc64a327acc, points: 0, grade: false, solution: true -->
```python
from math import sqrt
from cmath import sqrt as c_sqrt

def solve_quad(b, c):
    '''Solve the quadratic equation `x**2 + b*x + c == 0`. 
    
    Parameters
    ----------
    b, c : float
       Coefficients
       
    Returns
    -------
    x1, x2 : float or complex
        Roots.
    '''
    # YOUR CODE HERE
    D = (b**2) - 4.0 * c
    if 1 + c == b:
        return (-1.0), (-c)
    elif 1 + b + c == 0:
        return (1.0), (c)
    elif b == 0 and c == 0: 
        return 0, 0
    elif D**0.5 == b or D**0.5 == -b: 
        x1 = - b + c/b + c**2 / b**3 
        x2 = - (c/b + (c**2 / b**3))
        return x1, x2
    else:
        if D>=0: 
          #Formula Vieta
            x1 = - (2 * c) / (sqrt(D) + b)
            x2 = c / x2 
            return x1, x2
        elif D<0:
            x1 = - (2*c) / (c_sqrt(D) + b) 
            x2 = (2*c) / (c_sqrt(D) - b)
            return x1, x2
```

<!-- grade_id: cell-89cf7badb387b0e0, points: 5, grade: true, solution: false -->
```python
#
# Your implementation should pass tests in this cell.
#
# Do not remove or alter this cell. You may run it, but do not remove or alter it.
# Your changes to this cell will be ignored on grading.
# There will be additional tests.
#
from numpy.testing import assert_allclose
from numpy import allclose
variants = [{'b': 4.0, 'c': 3.0},
            {'b': 2.0, 'c': 1.0},
            {'b': 0.5, 'c': 4.0},
            {'b': 1e10, 'c': 3.0},
            {'b': -1e10, 'c': 3.0},
           ]

for var in variants:
    x1, x2 = solve_quad(**var)
    print(assert_allclose(x1*x2, var['c'], rtol=1e-15))
```

Output:
```
None
None
None
None
None
```

<!-- grade_id: cell-0970e4298cb77486, points: 5, grade: true, solution: false -->
```python
variants = [{'b': 4.0, 'c': 3.0},
            {'b': 2.0, 'c': 1.0},
            {'b': 0.5, 'c': 4.0},
            {'b': 1e10, 'c': 3.0},
            {'b': -1e10, 'c': 3.0},
           ]

for var in variants:
    x1, x2 = solve_quad(**var)
    print(allclose(x1*x2, var['c']))
```

Output:
```
True
True
True
True
True
```